# Module 1 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.

```python
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import sys
import math
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Creating a Python string

- Type "x = 'This is my first string' " without the outer quotation marks in the below cell and run it.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python

### YOUR CODE HERE
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
try:
    if ((x=='This is my first string') & (type(x)==str)):
        print('Great! You created a string!')
    else:
        print('x is not the right value check your code and run again')
except NameError:
    print('x is not defined! You need to run the code cell!')
    
assert_equal(type(x),str, msg='x is not a string!')
```

# Problem 2: Import numpy

- Type "import numpy" in the below cell to import the numpy module.


<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
if 'numpy' in dir():
    print('numpy module has been successfully imported')
else:
    print('numpy module has not been imported')
    
assert_equal(numpy.array([1])[0],1,msg='Make sure you ran the code cell to import numpy!')
```

# Problem 3: Add emphasis to Markdown text

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that does the following

    - Writes the word "bold" **in bold**
    
    - Writes the word "italics" *in italics*
    
    - Write the phrase "scrach and bold", first in italices, and then in bold, like so ~~**scratch and bold**~~.
    
    

- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the bold/italicize/scratch environments, and your text within those environments.

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
text = md.find('~~**')
assert_not_equal(text,-1,msg='Make sure you are using scratch first and bold second.')
text = md.find('**')
assert_not_equal(text,-1,msg='You did not bold the word bold.')
text = md.find('*ital')
assert_not_equal(text,-1,msg='You did not italicize the word italics')
```

# Problem 4: Create a Markdown list

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that creates a markdown list with 3 items. The items should be numbered 1 through 3, respectively, and there should be a `.` character after each number. Name each item "Item 1", "Item 2", and "Item 3", respectively. Note that the "I" in Item sould be capitalized.

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
text = md.find('1.')
assert_not_equal(text,-1,msg='Number 1. is mising from your markdown list')
text = md.find('2')
assert_not_equal(text,-1,msg='Number 2. is mising from your markdown list')
text = md.find('3.')
assert_not_equal(text,-1,msg='Number 3. is mising from your markdown list')
text = md.find('Item')
assert_not_equal(text,-1,msg='Make sure you named each item in the list properly.')
```

# Problem 5: Show code in Markdown

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that renders Python code. The Python code to be rendered is the statement `x = 'this is a pyhon string'`. 


- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the "code" environment.

<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
text = md.find('`x')
assert_not_equal(text,-1,msg='Make sure you are using the right characters to render the code.')
text = md.find('python')
assert_not_equal(text,-1,msg='Make sure you type the string correctly.')
```

-----

**&copy; 2018: Robert J. Brunner at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode